### Modeling Earth's Climate (Pfalzgraff, Eklof & Neshyba, 2026)
# If and For


## "If" blocks
Lots of times in climate modeling, we want to check whether some condition is met before going ahead. Python does this with an *if* block, which looks like

    if (some condition is met):
        line1
        line2
        ...

This is not actually Python code -- it's what we call *pseudo code* because it's supposed to be a little more intuitive than actual code$^*$. The point for now is, the entire set ("block") of indented commands line1, line2, ... will be executed only if that condition is met.

$*$That said, some people refer to Python as "executable pseudo code."


## "For" loops
A different situation arises when we want to do something over and over again, like march our climate model forward over time. That kind of thing calls for a *loop*. Among the many ways to set up a loop, the one we'll explore here is called a *for* loop. It looks a bit like the *if* syntax:

    for (each step in some sequence):
        line1
        line2
        ...

only now, those lines will be executed over and over again -- once for each step in the sequence. Our numerical method for simulating Earth's climate will be to set up a loop like this. We're going to specify exactly how far forward in time we want to go, and *for looping* is well suited for that kind of thing.

## Accumulating
As it turns out, inside our loop, we'll also want to collect values in an array, for plotting later. Sometimes we'll want to accumulate multiple arrays!


## Learning Goals
1. I'm familiar with the syntax of "if" blocks and "for" loops.
1. I can set up "for" loops of specified length.
1. I can accumulate values in a numpy array.
1. I can query the last few elements of an array, even if I don't know the length of the array.
1. I can plot multiple arrays on a single graph, and annotate using the "label/legend" method.
1. I can interpret exponential growth and decay curves in terms of $t_{1/2}$.

```python
%pip install -q ipympl
```

```python
%matplotlib widget
```

```python
# Bring in resources
import numpy as np
import matplotlib.pyplot as plt
```

### Syntax of an "if" statement
Below is an example of an if statement. You'll see that the test of whether two numbers are equal is a little unusual -- Python uses "==" instead of "=". The signs ">" and "<" work as expected, however. Check it out.

```python
print("testing whether 2+2<4")
if 2+2<4:
    print("This should not print because 2+2 is not less than 4")
    print("have a nice day")
    
print("testing whether 2+2=4")
if 2+2==4:
    print("This should print because 2+2 really does equal 4")
    print("have a nice day")
```

<!-- BEGIN QUESTION -->

### Your turn
In the cell below, set up three if-blocks like the ones above, but testing these conditions:

1. whether 2+2 is less than 3
1. whether 2+2 is equal to 3
1. whether 2+2 is greater than 3

Obviously, only the last should print stuff out.

_Points:_ 6

```python
...
```

<!-- END QUESTION -->

### A couple of examples of "for" loops
 We lay out a couple of examples in the cells below. Some notes about this:

- Like an "if" block, everything indented below the "for" line is part of the loop. Once you stop indenting, the loop is over. 
- We're using Python's "range" function to specify how many iterations we want. 
- The assumed lower limit of "range" is zero, and looping stops one short of the upper limit.

```python
# A loop that starts at 0 and goes to 4
for i in range(5):
    print(i)
print("i's last value was",i)
```

```python
# A loop that starts at 1 and goes to 4
for j in range(1,5):
    print(j)
print("j's last value was",j)
```

<!-- BEGIN QUESTION -->

### Your turn
Set up a loop using the "range" function that starts at 10 and goes to 13, printing the index as you go along.

_Points:_ 3

```python
...
```

<!-- END QUESTION -->

### Looping over a numpy array
This is a little different. We're going to create a numpy array first, and loop over it letting "range" figure out how long it is.

```python
# Here's a Numpy array
time = np.linspace(0,10,5)
print('The length of this array is', len(time))

# Here's looping over the contents of "time"
for i in range(len(time)):
    thistime = time[i]
    print(i,thistime)
```

### Your turn
Repeat almost exactly what you just did, but give *time* 9 elements, starting at 1 and ending at 20. Then loop over the elements in *time*, printing *i* and the time corresponding to it each iteration.

_Points:_ 3

```python
# Make a numpy array as described above, using np.linspace
...
# Loop over the elements in the "time" array and print i and the ith time
...
```

### Updating a variable inside a loop
Below is another example of a "for" loop. You can think of it as being in three parts:

- The lines *before* the loop, tagged as "Setting up", get things ready -- there's a variable (my_number) that's going to be getting bigger and bigger as the loop progresses.
- The loop itself. It starts with the word "for", which specifies how many times the indented lines below it will be executed. Each time through, my_number gets 5 added to it -- that's what "my_number = my_number + 5" does: it takes the old value of my_number, adds 5, and stores the result back in my_number.
- A reporting line that executes after the loop is completed. The reporting line isn't indented; if it were, it would be executed every iteration of the loop!

```python
# Setting up
my_number = 0

# Looping
for i in range(6):
    my_number = my_number + 5
    print(my_number)

# Reporting
print("After the loop, my_number=", my_number)
```

### Your turn
In the cell below, set up a "for" loop that does the same thing, but with some differences:
- Start off my_number at 10
- Increment my_number by 15 every time through the loop
- Go through the loop 4 times

_Points:_ 4

```python
...
```

### Accumulating
Sometimes, instead of just printing a variable to the screen as we go through a loop, we'd like to accumulate information as we go along. Here we'll do this with a *numpy array*.

We've already met numpy arrays -- that's what np.linspace makes. A numpy array can also be *empty*, which is how we set things up before the loop starts. Then, inside the loop, numbers are appended to it. The key commands in the cell below are

    my_number_array = np.empty(0)

before the loop, which creates an empty array, and

    my_number_array = np.append(my_number_array, my_number)

inside the loop, which tacks the current value of my_number onto the end of the array. Execute the cell below to see how this works. 

```python
# Setting up
my_number = 0

# Starting out with an empty numpy array
my_number_array = np.empty(0)

# Looping
for i in range(6):
    my_number = my_number + 5
    my_number_array = np.append(my_number_array, my_number)
    print(my_number_array)

# Reporting
print("After the loop, my_number_array=", my_number_array)
```

### Your turn
In the cell below, set up a "for" loop like the one above, but starting my_number at 0, incrementing it by 10, and going through the loop 10 times (so that my_number ends up at 100). Accumulate my_number in a numpy array named *my_number_array*.

_Points:_ 4

```python
...
```

### Examining elements of an array
When an array is short, it's fine to print it out in its entirety, as we've done. But if an array gets too long, there's no point in looking at all that. In such situations, it's often useful to examine the *first* and the *last* element of the array.

In the example below, we create a numpy array, called *my_array*, that contains five numbers. A few notes about this:

1. my_array[0] refers to the *first* element of this array.
1. my_array[4] refers to the *last* element of this array.
1. my_array[-1] refers to the last element of an array.
1. my_array[-2] refers to the second-to-last element of an array.

Why would we be interested in #3? It's just kind of handy -- it allows us to print the last element of an array without actually knowing how long it is. Here's an example:

```python
my_array = np.array([25, 32, 7, 12, 88])
print(my_array)
print("First element:", my_array[0])
print("Last element:", my_array[-1])
print("Second-to-last element:", my_array[-2])
```

<!-- BEGIN QUESTION -->

### Your turn
Using the above syntax, print the *first*, *last*, and *second-to-last* elements of the array you created a few cells back, which you named *my_number_array*.

_Points:_ 3

```python
...
```

<!-- END QUESTION -->

### Building arrays of numbers within a loop over time
Next we'll examine how to accumulate values *calculated* within a loop -- in this case, a loop over a time array, like the one we set up earlier. This time we'll accumulate two numpy arrays at once.

We're also taking this opportunity to introduce how to annotate two curves on the same graph using the "label/legend" method (which can be very handy). 

```python
# Create a time array (this one will have 30 elements)
time = np.linspace(0,20,30)

# Initialize a couple of empty numpy arrays
myarray1 = np.empty(0)
myarray2 = np.empty(0)

# Loop over the elements of time
for i in range(len(time)):
    
    # Extract the time in this iteration, and square it
    thisvalue1 = time[i]**2
    
    # Make another value that's twice what we just got
    thisvalue2 = 2*thisvalue1
    
    # append both of these values to our arrays
    myarray1 = np.append(myarray1,thisvalue1)
    myarray2 = np.append(myarray2,thisvalue2)

# Graph the results using the label/legend method
plt.figure()
plt.plot(time,myarray1,label='t^2')
plt.plot(time,myarray2,label='2xt^2')
plt.legend()
plt.grid(True)
```

### Your turn
Do this again, but with a new time array that has 500 elements ranging from 0 to 100 (years). In the loop, you'll construct two numpy arrays. One should describe exponential *growth*, which in mathematical terms would be written $e^{kt}$. We'll use $k=0.014$ for this. In Python, you'll need to write something like 

    thisvalue1 = np.exp(k*time[i])

We also want to compute exponential *decay*, i.e., $e^{-kt}$ (with the same value of $k$), for each time step:

    thisvalue2 = np.exp(-k*time[i])

Accumulate these values to their own numpy arrays (name them myarray1 and myarray2). Once the loop is done, plot these arrays on the same graph, and annotate them using the label/legend method.

_Points:_ 4

```python
# Define the rate constant
k = 0.014

# Create our time array
...
# Initialize a couple of empty numpy arrays using np.empty
...
# Loop over all the elements in time, calculate e^kt and e^(-kt), append them to your two numpy arrays
...
# Graph the results using the label/legend method; also turn the grid display to True
...
```

<!-- BEGIN QUESTION -->

### Pause for analysis
You may have noticed something about these two graphs: the time it took for $e^{kt}$ to double is about the same time it took for $e^{-kt}$ to drop by a factor of 2. We'll use the notation $t_{1/2}$ for this time.  Estimating from the graph, what's the value of $t_{1/2}$?

_Points:_ 3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Validating and finishing up
Assuming all this has gone smoothly, don't forget to do a Kernel/Restart & Run All, run the whole notebook, and make sure there aren't any errors.

Tests for accumulate (points: 4)
```
>>> assert len(my_number_array) == 10
```
failure_message: my_number_array should have 10 elements


Tests for exp_growth_decay (points: 4)
```
>>> assert len(myarray1) == 500 and len(myarray2) == 500
```
failure_message: myarray1 and myarray2 should each have 500 elements


Tests for time_array_loop (points: 3)
```
>>> assert len(time) == 9
```
failure_message: time should have 9 elements (np.linspace with 9 points)


Tests for update_variable (points: 4)
```
>>> assert my_number == 70
```
failure_message: After the loop, my_number should be 70 (start at 10, add 15 four times)
