### Modeling Earth's Climate (Pfalzgraff, Eklof & Neshyba, 2026)
# Simple Radiative Balance

## Absorbed Solar Radiation
The energy responsible for keeping Earth warm comes almost entirely from the Sun. This energy is all "light", but humans can see only part of it; of the rest, some comes in the form of ultraviolet light (which we can detect because it burns our skin) and near-infrared light (which is what you are feeling when you hold your hands out in front of a fire). All three forms of light (uv, visible, and near-IR) are collectively called *shortwave radiation*. 

How intense is shortwave radiation from the sun? To answer that, we need to talk about an instrument for measuring shortwave flux, called a *shortwave radiometer*. Shortwave radiometers measure the amount of shortwave energy (remember, that's uv, visible, and near-IR) passing through a 1 meter x 1 meter hoop. Real shortwave radiometers aren't actually that big, but they still tell you that key number: the energy intensity of sunlight on a per-meter-squared basis.) 

OK, now imagine you were are floating in space, just outside Earth's atmosphere, armed with your shortwave radiometer. It should read a number around $1367 \ W m^{-2}$. Is that a lot? Well, $1367 \ W m^{-2}$ is almost fourteen 100-Watt light bulbs, for every 1x1 square meter! This number is called the *solar constant*, and symbolized by $S_0$, so we say $S_0=1367 \ W m^{-2}$. $S_0$ varies a little over the course of a year (and of course is smaller for planets farther from the sun, bigger for planets that are closer), but $1367$ is a good average for us Earthlings.

As it turns out, however, that $1367 \ W m^{-2}$ doesn't actually make it down to every square meter of Earth's surface below you. There are two reasons for this:

1. The sun shines only on one side of the planet at a time, and it usually hits Earth's surface at an oblique angle. Turns out, this reduces the average energy potentially received by earth by a factor of 4, i.e., we're talking $S_0/4=342 \ Wm^{-2}$.
2. Clouds, snowpacks, desert sand, even forests, tend to reflect the shortwave light away, back out to space. The overall fraction reflected is called Earth's *albedo*, and it's about 30%. The albedo is typically given the symbol $\alpha$, and it's given as a fraction rather than a percent. So the statement "Earth's albedo is 30%" is written mathematically as $\alpha=0.3$. 

If we put this all together, we can write down something called the *Absorbed Shortwave Radiation*, which we'll symbolize here as $ASR$. It can be expressed as

$$
ASR = \dfrac{S_0}{4}(1-\alpha) \ \ \ \ (1)
$$

Like $S_0$, the units of $ASR$ are $W m^{-2}$. If you plug in the above numbers for $S_0$ and $\alpha$, you get $ASR=239 \ W m^{-2}$. That's still a lot if you think about it in terms of light bulbs: it's like there are over two 100-W lightbulbs of sunlight being absorbed by Earth's surface, day and night, over every square meter of the planet's surface.


## Outgoing Longwave Radiation
Imagine now that you are standing at the earth's surface, armed with a *longwave radiometer*. That's an instrument that measures the longwave radiant energy passing through a 1 meter x 1 meter hoop. To a good approximation, we'd expect that number to be given by the expression $\sigma T^4$, where $T$ is the temperature of the ground you're standing on. $\sigma$ is known as the Stefan-Boltzmann constant, $5.67 \times 10^{-8} W m^{-2}K^{-4}$. That's just physics, it doesn't depend on climate. 

For the purpose of radiative balance calculations, however, we need to re-position ourselves back to where we were at the beginning of this story: hovering in space just above Earth's atmosphere. Looking down at the earth's surface, the quantity our longwave radiometer measures now is called the Outgoing Longwave Radiation, or OLR. It will be, it turns out, not quite equal to the $\sigma T^4$ coming off Earth's surface. Why? Well first, something like 85% of the longwave radiation emitted by Earth's surface gets captured by water vapor, $CO_2$ and $CH_4$ in the atmosphere (they're all *greenhouse gases*). So we'd expect OLR to be reduced to something like 15% of $\sigma T^4$, right? However, it turns out that *the atmosphere itself emits longwave radiation*. Combining these influences, we can write 

$$ 
OLR = \kappa \sigma T^4 \ \ \ \ (2)
$$

where (as it turns out) $\kappa=0.614$. 

(Here's a more technical version of what we just said: Atmospheric greenhouse gases absorb 85% of the longwave photons that originate from Earth's surface, but those same gases emit their own longwave photons. Those longwave photons that are emitted by the atmosphere back down toward the earth's surface are responsible for greenhouse warming. Those longwave photons that are emitted by the atmosphere out to space combine with the 15% of Earth's surface photons that manage to get through the atmosphere unscathed, adding up to 61.4% of the orignal Stefan-Boltzmann ($\sigma T^4$) energy intensity coming off the earth's surface. That leads to the expression for the Outgoing Longwave Radiation in Eq. (2), with $\kappa=0.614$.)


## Radiative Balance
All this leads to an algorithm for modeling climate, based on radiative balance. It goes like this: suppose the earth were initially really cold. Then $T$ in Eq. (2) would be small, and that would mean the outgoing longwave radiation, $OLR$, would also be small -- smaller than our ASR of $239 \ W m^{-2}$ coming in. That means the earth's surface would be experiencing a *positive flux imbalance* -- more energy coming in than going out. What will the earth do then? Well, heat up of course! 

Things don't heat up instantly, however. If you're familiar with trying to boil water on a stovetop, you know that the more water in your pot, the longer it takes to heat up. We say that two gallons of water have twice the *heat capacity* as one gallon. The relevance of heat capacity to our climate model is, even when there's a positive flux imbalance, it would still take time to heat up. We can even write an equation for this:

$$
\Delta T = \dfrac{ASR-OLR}{C_p}\times \Delta t \ \ \ \ (3)
$$

where $\Delta T$ is the change in temperature of the earth's surface in the time $\Delta t$, and $C_P$ is our symbol for the heat capacity -- in our model, the heat capacity of a square meter of the average Earth's surface. It turns out that this second number is known: $C_P\approx 2.2\times 10^8 \ J \ m^{-2}K^{-1}$. And notice that it's in the denominator -- which means larger heat capacity means smaller change in temperature. Another term in this equation worth noticing is $ASR-OLR$. $ASR-OLR$ is likely to be positive when the earth is cold (because Eq. (2) says $OLR$ will be small), which means there would be a positive flux difference, which would mean $\Delta T>0$ (i.e. the temperature of the earth will *rise*). 

You can imagine the opposite process when Earth's temperature is initially very high. 

We haven't mentioned the factor $\Delta t$ yet. The lower case means we're talking about time, not temperature. And hopefully its appearance in Eq. (3) is fairly intuitive: the longer you heat something (bigger $\Delta t$), the more it will heat up (bigger $\Delta T$).

Equation (3) is the key equation in our model. It allows us to start with a given temperature, $T$, then march through time, at each time step either raising the temperature or lowering it, according to whether $ASR-OLR$ is positive or negative. In "pseudo-code", we could say we have a loop:

    for (each time step):
        calculate OLR from the current temperature using Eq. (2)
        calculate Delta_T from Eq. (3)
        update temperature using T(new) = T(old) + Delta_T

where Delta_T is what we've written here, mathematically, as $\Delta T$. 

How many time steps should we take? Well, there's no point in continuing on once there's no more change in temperature, right? In fact, what you'll see is that the temperature calculated in this way evolves toward a *steady-state* value: after enough time steps, Delta_T becomes so small that the temperature stops changing. If Earth is initially too cold, it gets warmer; if it's too hot, it cools down. That means it's not just any old steady state, it's a *stable* steady state. This is also called *homeostasis*.


## Climate sensitivity to changes in albedo
It turns out that in modeling the climate, a key thing one looks for is what's called the *climate sensitivity*. For example, we could say

$$
S = \dfrac {\Delta T_{SS}}{\Delta \alpha} \ \ \ \ (4)
$$

where $S$ is the sensitivity of Earth's temperature to its albedo ($\alpha$). To find this number from our model, we'd need to run the model once with a "reference" value of the albedo, and then again with a new value of the albedo. We'll call the difference in these two albedos $\Delta \alpha$, and the difference in the resulting steady-state temperatures  $\Delta T_{SS}$. From that, we can use Eq. (4) to calculate the climate sensitivity, $S$.

In the code below, you'll do this, and then use the resulting value of $S$ to predict the impact some new albedo might have on the temperature of the planet -- details to be given below.

## Caveats
The model described here ignores a key feature of climate called *feedbacks*. For example, if lots of snow melts because of higher temperatures (which it will do), Earth's albedo will go down, which will cause Earth's temperature to rise, which will cause more snow to melt: a positive feedback. But that's a challenge for another day.


## Learning Goals
1. I'm conversant with fundamental quantities associated with radiative balance, like $\alpha$, $S_o$, $ASR$, $OLR$, $\kappa$, and $\sigma$, and I can supply typical values, with units.
1. I can explain the sense in which the leveling-off of the temperature (and of the net flux) in the model is consistent with a stable steady-state temperature (or *homeostasis*).
1. I can explain how to calculate a *climate sensitivity (to albedo)*, and how to use that value as a predictive tool.

```python
%pip install -q ipympl
```

```python
%matplotlib widget
```

```python
import numpy as np
import matplotlib.pyplot as plt
```

### Model Definition
The parameters in this cell don't need to be changed from run to run (unless you want to change the model).

```python
def SRB(temperature_start,alpha):

    # Some constants
    solar_constant=342  # in units W m^-2
    kappa=0.614         # dimensionless
    sigma=5.67e-8       # in units W m^-2 K^-4
    heat_capacity=2.2e8 # in units J m^-2 K^-1

    # Time array (in years), and the time step
    time = np.linspace(0,30,301)
    delta_time_years = time[1]-time[0]           # Number of years for each time step
    delta_time = delta_time_years*24*60*60*365   # This converts the time step to seconds

    # Starting temperature
    temperature = temperature_start

    # Initializing the empty numpy arrays that will contain our results
    temperature_array = np.empty(0)
    fluxdiff_array = np.empty(0)

    # Looping over time
    for i in range(len(time)):

        # Outgoing longwave radiation (Wm^{-2})
        OLR=kappa*sigma*temperature**4

        # Absorbed solar radiation (also Wm^{-2})
        ASR=(1-alpha)*solar_constant

        # Flux difference (also Wm^{-2})
        fluxdiff=ASR-OLR

        # What's the temperature adjustment this time period?
        delta_temp = fluxdiff/heat_capacity*delta_time

        # Adjust Earth's temperature for this step
        temperature = temperature + delta_temp

        # Append to our arrays of temperatures and flux differences
        temperature_array = np.append(temperature_array,temperature)
        fluxdiff_array = np.append(fluxdiff_array,fluxdiff)

    return time, fluxdiff_array, temperature_array
```

<!-- BEGIN QUESTION -->

### Pause for analysis
Spend a little time getting to know this function:
1. The "for" loop runs for a fixed number of time steps. How many years does the model simulate, and how long (in years) is each time step?
1. There is one quantity that is calculated in the loop, that doesn't need to be in the loop, since it could easily (and more efficiently) be calculated before the loop begins. Which quantity is that?

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Experiment 1
Run the cell below to see how the model responds to an initial temperature of 200 degrees K and an albedo of 0.3.

```python
# Experiment 1

# Run the model
time_array_expt1, fluxdiff_array_expt1, temperature_array_expt1 = SRB(200,0.3)

# Plot the flux differences
plt.figure()
plt.plot(time_array_expt1,fluxdiff_array_expt1,'black')
plt.xlabel('time (years)')
plt.ylabel('Net Flux (W/m^2)')
plt.title('Experiment 1')
plt.grid(True)

# Plot the temperatures
plt.figure()
plt.plot(time_array_expt1,temperature_array_expt1,'black')
plt.xlabel('time (years)')
plt.ylabel('temperature (K)')
plt.title('Experiment 1')
plt.grid(True)

# Report the final (steady-state) temperature
print("The temperature ends up at ", temperature_array_expt1[-1])
```

<!-- BEGIN QUESTION -->

### Pause for analysis
In the cell below, write a few descriptive sentences about this:
1. Is the initial net flux positive or negative?
1. Is the temperature initially rising or falling?
1. What does it mean, physically, that the net flux levels off to a value of zero?

_Points:_ 3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Your turn
Carry out two more numerical experiments: 
- In the first cell below, carry out "Experiment 2". This experiment should have a starting temperature of 250 K. In your graph, use a blue line.
- In the following cell, carry out "Experiment 3". This experiment should have a starting temperature of 350 K. In your graph, use a red line.

Name your results the way Experiment 1 did, but with expt2 or expt3 in place of expt1 -- so, for example, your temperatures should end up in *temperature_array_expt2* and *temperature_array_expt3*.

_Points:_ 5

```python
# Experiment 2

# Run the model with a starting temperature of 250 K
...
# Plot the flux differences
...
# Plot the temperatures
...
# Report the final (steady-state) temperature
...
```

```python
# Experiment 3

# Run the model with a starting temperature of 350 K
...
# Plot the flux differences
...
# Plot the temperatures
...
# Report the final (steady-state) temperature
...
```

<!-- BEGIN QUESTION -->

### Your turn (again)
Assuming you've given your results unique names, plot all three temperatures on the same graph, as a function of time.

_Points:_ 4

```python
...
```

<!-- END QUESTION -->

### Climate sensitivity
Here, we'd like to see how the steady-state temperature depends on the albedo. In order to calculate it, we could go about it in the following way. First, specify a new value of the albedo. For example, you could set up Experiment #4, with $\alpha=0.35$. Then, since Experiment #3 used $\alpha=0.3$, we could say

$$
\Delta \alpha = 0.05
$$

When you run Experiment #4, you'll get a different steady-state temperature, of course. If we said $T_4$ is that temperature, and $T_3$ was the steady-state temperature you got in experiment #3, you could find the change in temperature resulting from the change in albedo by calculating

$$
\Delta T_{SS} = T_4-T_3
$$

Of course, the bigger the change in the albedo, the bigger we expect $\Delta T$ to be. What's key, therefore, is the ratio, $\dfrac {\Delta T_{SS}}{\Delta \alpha}$. As described in the introduction, this is called the *climate sensitivity* to changes in albedo. Doing this in Python, it will be helpful to create a variable holding this value, with a line like the following:

    sensitivity = delta_T_SS/delta_alpha

Note that the sign of the sensitivity is important here: if you get a negative number, you're probably on track. 

_Points:_ 3

```python
...
```

### Pause for analysis
But *will* Earth's albedo change in a warming climate? That's a big question, which we won't be able to do justice to here. What we can do, is use the foregoing Simple Radiative Balance ("SRB") theory to predict the temperature impact of a hypothetical change. Here's how: if you solve Eq. (4) for $\Delta T_{SS}$, you get 

$$
\Delta T_{SS} = S \ \Delta \alpha
$$

So what's a reasonable guess for $\Delta \alpha$? A recent study by Goode et al, described in a press release at https://news.agu.org/press-release/Earth-is-dimming-due-to-climate-change, reported that over the past two decades, Earth's average albedo dropped from $0.30$ to $0.2985$ -- a change in albedo of $-0.0015$. In that case, the corresponding Python code could look like

    deltaT_SS = sensitivity*(-0.0015)
    
In the cell below, use SRB theory to predict the change in temperature resulting from this reduction in albedo. 

_Points:_ 2

```python
...
```

<!-- BEGIN QUESTION -->

### Pause for analysis

Read the short article/press release linked above, then answer the following questions:

1. What phenomenon are Goode et al actually measuring, that led to the reduction in albedo they report? It will help to take a look at the press release and the paper by Goode et al to answer this. Hint: the phenomenon has been noticed for a long time -- including by Leonardo da Vinci!
1. By comparison with measurements by NASA’s Clouds and the Earth’s Radiant Energy System, Goode et al speculate a reason behind the albedo reduction. What is that reason?

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Validating and finishing up
Assuming all this has gone smoothly, don't forget to do a Kernel/Restart & Run All, run the whole notebook, and make sure there aren't any errors.

Tests for albedo_impact (points: 2)
```
>>> assert bool(0.1 < deltaT_SS < 0.25)
```
failure_message: deltaT_SS should be a small positive warming, around 0.16 K


Tests for expts_2_and_3 (points: 5)
```
>>> assert bool(np.isclose(temperature_array_expt2[-1], 288, atol=0.1))
```
failure_message: Experiment 2 should end near the steady-state temperature, about 288 K
```
>>> assert bool(np.isclose(temperature_array_expt3[-1], 288, atol=0.1))
```
failure_message: Experiment 3 should end near the steady-state temperature, about 288 K


Tests for sensitivity (points: 3)
```
>>> assert bool(-120 < sensitivity < -90)
```
failure_message: sensitivity should be roughly -100 K per unit of albedo (note the sign)
